<p style="font-family: Arial; font-size:3.75vw;color:purple; font-style:bold"><br>
Classification Exercise Notebook
</p><br>

# Exercise Notebook Instructions

### 1. Important: Only modify the cells which instruct you to modify them - leave "do not modify" cells alone.  

The code which tests your responses assumes you have run the startup/read-only code exactly.

### 2. Work through the notebook in order.

Some of the steps depend on previous, so you'll want to move through the notebook in order.

### 3. It is okay to use numpy libraries.

You may find some of these questions are fairly straightforward to answer using built-in numpy functions.  That's totally okay - part of the point of these exercises is to familiarize you with the commonly used numpy functions.

### 4. Seek help if stuck

If you get stuck, don't worry!  You can either review the videos/notebooks from this week, ask in the course forums, or look to the solutions for the correct answer.  BUT, be careful about looking to the solutions too quickly.  Struggling to get the right answer is an important part of the learning process.

```python
# DO NOT MODIFY

# import appropriate libraries
import pandas as pd
import numpy as np
from sklearn.model_selection import train_test_split
from sklearn.metrics import accuracy_score
from sklearn.tree import DecisionTreeClassifier
from sklearn.linear_model import LinearRegression
from sklearn.tree import DecisionTreeRegressor
from sklearn.metrics import mean_squared_error
from math import sqrt
weather_path = '../../../datasets/weather/daily_weather.csv' 
```

```python
# DO NOT MODIFY

# we will use the daily weather dataset of this exercise.

def get_data():
    return pd.read_csv(weather_path)

df = get_data()
```

<p style="font-family: Arial; font-size:2.75vw;color:purple; font-style:bold"><br>

Exercise 1: Detect if NULLs Exist in the Data<br><br></p>


In the cell below, modify the function to RETURN True if there exists a NULL value in data given as input, otherwise False

<!-- grade_id: c555, points: 0, grade: false, solution: true -->
```python
# modify this cell

def detect_nulls(matx):
    ### BEGIN SOLUTION
    return matx.isnull().any().any()
    ### END SOLUTION
```

<!-- grade_id: ex15, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY
ans = True

try: assert np.alltrue(detect_nulls(df) == ans)
except AssertionError as e: print("Try again, your output did not match the expected answer above")
```

<p style="font-family: Arial; font-size:2.75vw;color:purple; font-style:bold"><br>

Exercise 2: Data Cleaning<br><br></p>


In the cell below, modify the function to remove every row that has atleast one NULL,
and RETURN a new data frame.

<!-- grade_id: 5551, points: 0, grade: false, solution: true -->
```python
# modify this cell

def clean_data(matx):
    ### BEGIN SOLUTION
    matx = matx.dropna() # creates a slice of the data
    return matx
    ### END SOLUTION
```

<!-- grade_id: ex151, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY

try: 
    df2=clean_data(df)
    assert np.alltrue(df2.shape[0] == 1064)
except AssertionError as e: print("Try again")
```

```python
df2.describe()
```

Output:
```
number  air_pressure_9am  air_temp_9am  avg_wind_direction_9am  \
count  1064.000000       1064.000000   1064.000000             1064.000000   
mean    545.001880        918.903180     65.022609              142.306756   
std     316.796111          3.179040     11.168033               69.149472   
min       0.000000        907.990000     36.752000               15.500000   
25%     269.750000        916.595376     57.398000               65.979244   
50%     543.500000        918.942281     65.778479              165.937461   
75%     819.250000        921.169054     73.530872              191.100000   
max    1094.000000        929.320000     98.906000              343.400000   

       avg_wind_speed_9am  max_wind_direction_9am  max_wind_speed_9am  \
count         1064.000000             1064.000000         1064.000000   
mean             5.485793              148.480424            6.999714   
std              4.534427               67.154911            5.590790   
min              
```

<p style="font-family: Arial; font-size:2.75vw;color:purple; font-style:bold"><br>

Exercise 3: Convert to a Classification Task<br><br></p>

In the cell below, modify the function that takes three inputs: a dataframe, a column name C, and a threshold value X.
The function has to create a new column called 'target' which is 1
when given column C is > X, otherwise 0. Return a new dataframe with this new column, and keep all the 
previous columns.

<!-- grade_id: c5552, points: 0, grade: false, solution: true -->
```python
# modify this cell

def classification_task(df, C, X):
    ### BEGIN SOLUTION
    clean_data = df.copy()
    clean_data['target'] = (clean_data[C] > X)*1
    return clean_data
    ### END SOLUTION
```

<!-- grade_id: ex152, points: 1, grade: true, solution: false -->
```python

# DO NOT MODIFY

ans = 558

try: 
    c = classification_task(df2, 'air_temp_9am', 65.0)
    assert np.alltrue(c['target'].sum() == ans)
except AssertionError as e: print("Try again, your solution did not produce the expected output above")
```

```python
del c['number']
```

```python
def get_ready(df):
    from sklearn.model_selection import train_test_split
    # declare features to use
    # use the 9am sensor signals to predict the 9am air temperature
    features = ['air_pressure_9am','avg_wind_direction_9am','avg_wind_speed_9am','max_wind_direction_9am']
    # create copy
    df_copy = df.copy()
    df_copy['target'] = (df_copy['air_temp_9am'] > 65.0)*1
    # separate features from target
    X=df_copy[features]
    y=df_copy['target']
    # perform split and return
    return train_test_split(X, y, test_size=0.33, random_state=983)
```

```python
X_train, X_test, y_train, y_test = get_ready(pd.read_csv(weather_path).dropna())
```

<p style="font-family: Arial; font-size:2.75vw;color:purple; font-style:bold"><br>

Exercise 4: Create a Classifier<br><br></p>

In the cell below, modify the function to take X_train, y_train only and RETURN a classifier
for predicting the y_train based on columns in X_train. You can pick any classifier model.

The function should RETURN a trained model. We will test your classifier on X_test and y_test

<!-- grade_id: c5553, points: 0, grade: false, solution: true -->
```python
# modify this cell

def train_classifier(X_train, y_train):
    ### BEGIN SOLUTION
    temperature_classifier = DecisionTreeClassifier(max_leaf_nodes=45, random_state=3)
    temperature_classifier.fit(X_train, y_train)
    return temperature_classifier
    ### END SOLUTION
```

<!-- grade_id: ex153, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY

min_score = 0.63

try: 
    v = train_classifier(X_train, y_train)
    y_pred = v.predict(X_test)
    print(accuracy_score(y_test, y_pred))
    assert np.alltrue(accuracy_score(y_test, y_pred) > min_score)
except AssertionError as e: print("Keep trying - can you get an accuracy_score > %f" % min_score)
```

Output:
```
0.6363636363636364
```

```python
# Lets try it again with a different model, like a regression model
X_train, X_test, y_train, y_test = get_ready(pd.read_csv(weather_path).dropna())
regressor = LinearRegression()
regressor.fit(X_train,y_train)

```

Output:
```
LinearRegression(copy_X=True, fit_intercept=True, n_jobs=None,
         normalize=False)
```

```python
y_prediction = regressor.predict(X_test)
y_prediction[:20]%20
```

Output:
```
array([ 0.46615388,  0.57847173,  0.68728384,  0.51966368,  0.48615424,
        0.47485446, 19.94412973,  0.354388  ,  0.13548431,  0.45932973,
        0.46321121,  0.04297255,  0.41060332,  0.25587705,  0.65101399,
        0.56185222,  0.49563279,  0.58075645,  0.6330126 ,  0.43193377])
```

```python
y_test.describe()
```

Output:
```
count    352.000000
mean       0.590909
std        0.492366
min        0.000000
25%        0.000000
50%        1.000000
75%        1.000000
max        1.000000
Name: target, dtype: float64
```

```python
RMSE = sqrt(mean_squared_error(y_true = y_test, y_pred = y_prediction))
print(RMSE)
```

Output:
```
0.4826991986559557
```

```python

```

Output:
```
array([ 0.46615388,  0.57847173,  0.68728384,  0.51966368,  0.48615424,
        0.47485446, -0.05587027,  0.354388  ,  0.13548431,  0.45932973,
        0.46321121,  0.04297255,  0.41060332,  0.25587705,  0.65101399,
        0.56185222,  0.49563279,  0.58075645,  0.6330126 ,  0.43193377,
        0.46487578,  0.36677358,  0.48374542,  0.62526081,  0.54166881,
        0.56194202,  0.56445309,  0.5620413 ,  0.44434999,  0.66370137,
        0.47257176,  0.5540992 ,  0.586924  ,  0.49611141,  0.62433763,
        0.49485149,  0.67629367,  0.35741511,  0.33022921,  0.53453309,
        0.12770079,  0.53108975,  0.68458056,  0.60008073,  0.56596215,
        0.67015761,  0.65207003,  0.66797241,  0.57262911,  0.15576615,
        0.42666212,  0.45454543,  0.43458598,  0.54822863,  0.57273664,
        0.57845712,  0.70754532,  0.60667636,  0.49393397,  0.37061067,
        0.59284115,  0.6075016 ,  0.59080294,  0.52859403,  0.49998573,
        0.47075324,  0.6483956 ,  0.40531934,  0.44853561,  0.54
```

```python

```

```python

```

Output:
```
0.4826991986559557
```

```python

```